
# Assignment No 4b
###### *Sibt ul Hussain*
----
## Goal

Your goal in this part of assigment is to implement a Decision Tree Classifier for categorical variables.

**Note** Please note that you are allowed to use only those libraries which we have discussed in the class, i.e. numpy, scipy, pandas.

## Submission Instructions
You are required to submit the original notebook file on the Slate (with .ipynb extension), with complete set of outputs. Students failing to do so will get zero marks. 

*Please read each step carefully and understand it fully before proceeding with code writing*

## Plagiarism
Any form of plagiarism will not be tolerated and result in 0 marks.



### Decision Tree Classifier

Now in this assignment we will be implementing the Decision Classifier for both Continuous and Categorical attributes.

Decision tree can be built by using any of the following split criterias, namely:
 - Information Gain
 - Gini Index
 - CART 

However, you are required here to implement the decision tree with information gain as splitting criterion.

Remember in my code i am not looking for maximizing the information gain, instead i am looking for minimizing the split entropy. Recall,
$$Information Gain  = H(D) - H(D_Y,D_N)$$

Where,

$H(D)$ is the data set entroy and $H(D_Y,D_N)$ is split entropy. Since $H(D)$ is constant for the given dataset so maximizing the entropy is equal to minimizing the split entropy and that is what is being represented in my code outputs.

```python
%pylab inline
import scipy.stats
from collections import defaultdict  # default dictionary 
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

```python
#Cutomize the Matplotlib for beautiful plots...
# import dmStyle
# dmStyle.customize_mpl()

```

```python
import pandas as pd
import tools as t # set of tools for plotting, data splitting, etc..
```

```python
def getSplits(categories):
    '''
        function returns list of splits for the given list of categorical variables...
        
        Input:
        ------------
            categories: a list of unique categories...
        
        Return:
        ------------
            list of splits(tuples) for given list of categorical variables. Each pair of sublists
            defines the left and right splits, e.g. This list
            [('y', 'f'), ('s', 'g'), ('f'), ('y', 's', 'g')]
            
            defines two splits with each pair representing a different split.
        Examples:
        ------------
        splits=getSplits(['a1','a2','a3','a4']) will return 
        [('a1', 'a2', 'a4'), ('a3',), ('a2', 'a4'), 
        ('a1', 'a3'), ('a1', 'a4'), ('a3', 'a2'), ('a1', 'a3', 'a2'), 
        ('a4',), ('a3', 'a4'), ('a1', 'a2'), ('a1',), ('a3', 'a2', 'a4'), ('a2',), ('a1', 'a3', 'a4')]

            
    '''
    categories=set(categories)
    tsplits=t.get_powerset(categories,len(categories)-1) # get all the power sets with the given cardinality...
    flist=[]
    return tsplits

```

```python
def listOfSplits(Unique,Set):
    
    if (len(Set) > 1):
        splits = []
        length = Unique.shape[0] - 1
        size = len(Set)
        for i in range(0,size):
            s = Set.pop()
            if (len(s) != length):
                splits.append(np.array(list(s)))

        return np.array(splits)
    else:
        return tuple(Set)
```

```python
splits=getSplits(['a1','a2','a3','a4'])
print splits,len(splits)
# print listOfSplits(np.array(['a1','a2','a3','a4']),splits)
```

Output:
```
set([frozenset(['a1', 'a2', 'a4']), frozenset(['a2', 'a4']), frozenset(['a1', 'a4']), frozenset(['a3']), frozenset(['a1', 'a3', 'a2']), frozenset(['a3', 'a2']), frozenset(['a4']), frozenset(['a1', 'a3']), frozenset(['a3', 'a4']), frozenset(['a1']), frozenset(['a3', 'a2', 'a4']), frozenset(['a1', 'a2']), frozenset(['a2']), frozenset(['a1', 'a3', 'a4'])]) 14
```

```python
class Node:
    def __init__(self,purity,klasslabel='',score=0,split=[],fidx=-1):
        '''
            purity: purity level at which to stop
            klasslabel: klasslabel of the node, (for leaf node)
            score: information gain of the newly added node
            split: splitting threshold
            fidx: feature index            
        '''
        self.lchild=None
        self.rchild=None
        self.klasslabel=klasslabel        
        self.split=split
        self.score=score
        self.fidx=fidx
        self.purity=purity
        self.ftype= 'categorical' if type(self.split) in [tuple, str, numpy.string_] else 'continuous' # feature type 

        
    def set_childs(self,lchild,rchild):
        # YOUR CODE HERE
        self.lchild = lchild
        self.rchild = rchild
        pass

        
    def isleaf(self):
        # YOUR CODE HERE
        return (self.lchild == None and self.rchild == None)
        pass

        
    def isless_than_eq(self, X):
        # YOUR CODE HERE
        pass

        
    def get_str(self):        
        if self.isleaf():
            return 'C(class={},Purity={})'.format(self.klasslabel,self.purity)
        else:
            return 'I(Fidx={},Score={},Split={})'.format(self.fidx,self.score,self.split)  
```

```python
# A placeholder class 
# TODO: You have to implement the following class, remember from the lectures that you will 
# need to build a model for each different class you are trying to identify...
# dt=DecisionTree(0.95,5,5)
# dt.evaluate_categorical_attribute(tx,Y)
a = np.array(['a2', 'a4'])
b = np.array(['a1' ,'a2' ,'a3' ,'a4'])
cc = []
cc.append(np.where(b==a[0])[0])
cc.append(np.where(b==a[1])[0])
print np.array([]).shape[0]
# np.setdiff1d(b,a)
```

Output:
```
0
```

```python
### import pdb
## Your code goes here...
# You might need to define auxliary classes for composition.. ?
class DecisionTree:
    ''' Implements the Decision Tree For Classification... '''
    def __init__(self, purityp, exthreshold, maxdepth=10):        
        self.purity = purityp
        self.exthreshold = exthreshold
        self.maxdepth = maxdepth
        self.tree = None
        
    def train(self, X, Y):
        ''' Train Decision Tree using the given 
            X [m x d] data matrix and Y labels matrix
            
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns:
            -----------
            Nothing
            '''
        nexamples, nfeatures = X.shape
        # # now go and train a model for each class...
        # YOUR CODE HERE
        self.tree = self.build_tree(X,Y,self.maxdepth)
        
        
    def build_tree(self, X, Y, depth):
        """ 
            Function is used to recursively build the decision Tree 
          
            Input
            -----
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns
            -------
            root node of the built tree...
        """
        nexamples, nfeatures = X.shape
        
        klasses=np.unique(Y);
        # YOUR CODE HERE
        
        size = X.shape[0]
        unique,counts = np.unique(Y,return_counts=True)
        purity = 0
        print "size is ",size
        if (size < 0):
            return
        if (counts.shape[0] > 0):
            purity = max(counts)/float(sum(counts))
#         if ((depth <= 0 and size < self.exthreshold) or (purity >= self.purity and depth <= 0)):
        if (size <= self.exthreshold or purity >= self.purity) or depth <= 0:
            c = ''
            if (counts.shape[0] > 0):
                c = unique[np.argmax(counts)]
                print "c is ",c
                return Node(purity,c)
            leafNode = Node(purity,c)
            return leafNode
        V = []
        S = []
        xlid = []
        xrid = []
        
        for i in range(0,X.shape[1]):
            v,s,xlidx,xridx = self.evaluate_categorical_attribute(X[:,i],Y)
            
            V.append(v)
            S.append(s)
            xlid.append(xlidx)
            xrid.append(xridx)
        S = np.array(S)
        ind = np.argmin(S)
        node = Node(purity)
        node.score = S[ind]
        node.fidx = ind
        node.split = V[ind]
        Dy = X[xlid[ind]]
        Dn = X[xrid[ind]]
        Dyl = Y[xlid[ind]]  # Dyl Labels
        Dnl = Y[xrid[ind]]  # Dnl labels
#         print Dy,Dyl
#         print Dn,Dnl
#         print "Depth is ", depth
#         node.set_childs(self.build_tree(Dy,Dyl,depth-1),self.build_tree(Dn,Dnl,depth-1))
        if Dy.shape[0] <= 0 or Dn.shape[0] <= 0:
            return Node(purity,unique[np.argmax(counts)])
        node.lchild = self.build_tree(Dy,Dyl,depth-1)
        node.rchild = self.build_tree(Dn,Dnl,depth-1)
#         self.tree = node
#         print "arbaz"
        return node

        
        
        
        pass
                
    def test(self, X):
        
        ''' Test the trained classifiers on the given set of examples 
        
                   
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional test examples.
           
            Returns:
            -----------
                pclass: the predicted class for each example, i.e. to which it belongs
        '''
        
        nexamples, nfeatures = X.shape
        pclasses = self.predict(X)
        
        # your code go here...
    def calculateEntropy(self,row):
        e = 0
        total = sum(row)
#         print row,"values"
#         print "sss",row.shape
        for i in range(0,row.shape[0]):
            temp = row[i]/float(total)
            if temp == 0:
                temp = 1
            e = e + (row[i]/float(total))*np.log2(temp)
        e = -e
        return e
        
        pass
            
    def calculateSplitScoreCategorical(self,combinationDy,unique,matrix):
        # unique = unique values of featuere; matrix = [m x k] size
        dyTable = []
        dnTable = []
        dd = np.array([])
        for i in range(0,len(combinationDy)):
            dyTable.append(matrix[(combinationDy[i] == unique)][0])
#             print matrix[(combinationDy[i] == unique)][0],"yeh ha"
#             np.concatenate((dd,matrix[(combinationDy[i] == unique)]))
        combDn = np.setdiff1d(unique,combinationDy)
        for j in range(0,combDn.shape[0]):
            dnTable.append(matrix[(combDn[j] == unique)][0])
#         dyTable = matrix[(combination == unique)]
#         dnTable = matrix[(combination != unique)]
        dyTable = np.array(dyTable)
        dnTable = np.array(dnTable)
#         print matrix,combinationDy,"Table:",dyTable,dyTable.shape
#         print matrix,combinationDy,"Table:",dnTable
        dySum = np.sum(dyTable,axis=0)
        dnSum = np.sum(dnTable,axis=0)
#         print self.purity
#         print "MAtrix",matrix
#         print "Dy table",dyTable
#         print "Dn tables",dnTable
#         print dySum,"y"
#         print dnSum,"n"
        hDy = 100
        hDn = 100
        if (dyTable.shape[0] != 0):
            hDy = self.calculateEntropy(dySum)
        if (dnTable.shape[0] != 0):
            hDn = self.calculateEntropy(dnSum)
#         print (sum(dySum)/float(sum(dySum)+sum(dnSum)))*hDy + (sum(dnSum)/float(sum(dySum)+sum(dnSum)))*hDn
#         print dySum.sum()
#         print (((dySum.sum())/float((dySum)+(dnSum.sum())))*hDy + ((dnSum.sum())/float((dySum.sum())+(dnSum.sum())))*hDn)
        return (sum(dySum)/float(sum(dySum)+sum(dnSum)))*hDy + (sum(dnSum)/float(sum(dySum)+sum(dnSum)))*hDn
        

        pass
    
    def evaluate_categorical_attribute(self, feat, Y):
        """
            Evaluates the cateogrical attribute for all possible split points, i.e. 2^(m-1)-2 for
            possible feature selection

            Input:
            ---------
            feat: a categorical feature
            Y: labels

            Returns:
            ----------
            v: splitting threshold
            score: splitting score
            Xlidx: Index of examples belonging to left child node
            Xridx: Index of examples belonging to right child node

        """
        categories = set(feat)
        splits = getSplits(categories) if len(categories) > 1 else tuple(categories)
        freq = scipy.stats.itemfreq(Y)

        # YOUR CODE HERE
        unique = np.unique(feat)
        classes = np.unique(Y)
        totalClasses = len(unique)
        N = np.zeros((totalClasses,freq.shape[0]))  # matrix of size m x k; m = number of unique values of feature;
                                                    # k = total number of classes
        for i in range(0,feat.shape[0]):
            N[np.where(feat[i] == unique)[0][0],np.where(Y[i] == classes)[0][0]] += 1 
#         print N
#         print freq
        splits = listOfSplits(unique,splits)
        length = len(splits)
        splitScore = []
        for i in range(0,length):
            splitScore.append(self.calculateSplitScoreCategorical(splits[i],unique,N))
        score = 100
        split = ''
        xlidx = []
        xridx = []
        if (len(splitScore) > 0):
            ind = np.nanargmin(splitScore)
            score = splitScore[ind]
            split = splits[ind]
            splitDn = np.setdiff1d(unique,split)
            xlidx = np.array([])
            xridx = np.array([])
            for i in split:
                xlidx = np.concatenate((xlidx,np.where(feat==i)[0]))
            for j in splitDn:
                xridx = np.concatenate((xridx,np.where(feat==j)[0]))
            xlidx = xlidx.astype(int)
            xridx = xridx.astype(int)
#         print split,score
#         print Y[xlidx]
#         print Y[xridx]
        return split,score,xlidx,xridx
        
            
            
            
        

        
    def evaluate_numerical_attribute(self, feat, Y):
        '''
            Evaluates the numerical attribute for all possible split points for
            possible feature selection
            
            Input:
            ---------
            feat: a contiuous feature
            Y: labels
            
            Returns:
            ----------
            v: splitting threshold
            score: splitting score
            Xlidx: Index of examples belonging to left child node
            Xridx: Index of examples belonging to right child node
            
        '''
        
        #classes = np.unique(Y)
        #nclasses = len(classes)
        #f = np.sort(feat)
        
        # YOUR CODE HERE
#         raise NotImplementedError()
        
        
        
        classes=np.unique(Y)
        nclasses=len(classes)
        sidx=np.argsort(feat)
        f=feat[sidx] # sorted features
        sY=Y[sidx] # sorted features class labels...
        
        # YOUR CODE HERE
                    

                        
    def predict(self, X):
        
        """
        Test the trained classifiers on the given example X
        
                   
            Input:
            ------
            X: [1 x d] a d-dimensional test example.
           
            Returns:
            -----------
                pclass: the predicted class for the given example, i.e. to which it belongs
        """
        return self._predict(self.tree,X)

    
    def _predict(self, node, X):
        # YOUR CODE HERE
        pclasses = []
        for i in X:
            temp = node
            while not(temp.isleaf()):
                count = 0
                for j in temp.split:
                    if (i[temp.fidx] == j):
                        count += 1
                if (count > 0):
                    temp = temp.lchild
                else:
                    temp = temp.rchild
#                 print "arbaz"
            pclasses.append(temp.klasslabel)
        return pclasses
        pass
        


    def __str__(self):
        str = '---------------------------------------------------'
        str += '\n A Decision Tree With Depth={}'.format(self.find_depth())
        str += self.__print(self.tree)
        str += '\n---------------------------------------------------'
        return str  # self.__print(self.tree)        
        
    def find_depth(self):
        return self._find_depth(self.tree)
    
    def _find_depth(self, node):
        if not node:
            return
        if node.isleaf():
            return 1
        else:
            return max(self._find_depth(node.lchild), self._find_depth(node.rchild)) + 1
    def __print(self, node, depth=0):
        
        ret = ""

        # Print right branch
        if node.rchild:
            ret += self.__print(node.rchild, depth + 1)

        # Print own value
        
        ret += "\n" + ("    "*depth) + node.get_str()

        # Print left branch
        if node.lchild:
            ret += self.__print(node.lchild, depth + 1)
        
        return ret
```

```python
dt=DecisionTree(0.95,5,5)
dt.evaluate_categorical_attribute(tx,Y)
print np.unique(Y,return_counts=True)
# a = np.array([1,2,3,4,5,6])
# a = a.reshape((3,2))
# print a
# print np.sum(a,axis=1)
```

Output:
```
(array(['Iris-setosa', 'Iris-versicolor'], dtype=object), array([ 50, 100]))
```

# Lets test our code for the given example in the book.

```python
#load the Iris dataset
tdata=pd.read_csv('./iris.data', header=None)
tdata.columns=['SepalLength','SepalWidth','PetalLength','PetalWidth','Class']
```

```python
tx=tdata['SepalLength'].dropna()
tx[(tdata['SepalLength']>=4.3) & (tdata['SepalLength']<=5.2)]='a1'
tx[(tdata['SepalLength']>5.2) & (tdata['SepalLength']<=6.1)]='a2'
tx[(tdata['SepalLength']>6.1) & (tdata['SepalLength']<=7.0)]='a3'
tx[(tdata['SepalLength']>7.0) & (tdata['SepalLength']<=7.9)]='a4'
```

```python
print tx.values
Y=tdata['Class'].dropna()
Y[Y=='Iris-virginica']='Iris-versicolor'
Y=Y.values
print Y
```

Output:
```
['a1' 'a1' 'a1' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a2'
 'a2' 'a2' 'a1' 'a2' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1'
 'a1' 'a2' 'a1' 'a2' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1'
 'a1' 'a1' 'a1' 'a2' 'a1' 'a3' 'a3' 'a3' 'a2' 'a3' 'a2' 'a3' 'a1' 'a3' 'a1'
 'a1' 'a2' 'a2' 'a2' 'a2' 'a3' 'a2' 'a2' 'a3' 'a2' 'a2' 'a2' 'a3' 'a2' 'a3'
 'a3' 'a3' 'a3' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2' 'a3' 'a3' 'a2' 'a2'
 'a2' 'a2' 'a2' 'a1' 'a2' 'a2' 'a2' 'a3' 'a1' 'a2' 'a3' 'a2' 'a4' 'a3' 'a3'
 'a4' 'a1' 'a4' 'a3' 'a4' 'a3' 'a3' 'a3' 'a2' 'a2' 'a3' 'a3' 'a4' 'a4' 'a2'
 'a3' 'a2' 'a4' 'a3' 'a3' 'a4' 'a3' 'a2' 'a3' 'a4' 'a4' 'a4' 'a3' 'a3' 'a2'
 'a4' 'a3' 'a3' 'a2' 'a3' 'a3' 'a3' 'a2' 'a3' 'a3' 'a3' 'a3' 'a3' 'a3' 'a2']
['Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setos
```

```python
%pdb off
dt=DecisionTree(0.95,5,5)
split, gain, Xlidx, Xridx = dt.evaluate_categorical_attribute(tx,Y)
print 
print split, gain
```

Output:
```
Automatic pdb calling has been turned OFF

['a1'] 0.508690715247
```

<!-- grade_id: split_gain, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal, assert_almost_equals

%pdb off
dt=DecisionTree(0.95,5,5)
split, gain, Xlidx, Xridx = dt.evaluate_categorical_attribute(tx,Y)


assert_almost_equal('a1', split[0])
assert_almost_equal(gain, 0.51, 1)
```

Output:
```
Automatic pdb calling has been turned OFF
```

# Now lets test on [Mushroom](https://archive.ics.uci.edu/ml/machine-learning-databases/mushroom/agaricus-lepiota.names) dataset


This data set includes descriptions of hypothetical samples corresponding to 23 species of gilled mushrooms in the Agaricus and Lepiota Family (pp. 500-525).  Each species is identified as definitely edible, definitely poisonous, or of unknown edibility and not recommended.  This latter class was combined with the poisonous one.  The Guide clearly states that there is no simple rule for determining the edibility of a mushroom; no rule like ``leaflets three, let it be'' for Poisonous Oak and Ivy.


- Number of Instances: 8124

- Number of Attributes: 22 (all nominally valued)

- Attribute Information: (classes: edible=e, poisonous=p)
         1. cap-shape:                bell=b,conical=c,convex=x,flat=f, knobbed=k,sunken=s
         2. cap-surface:              fibrous=f,grooves=g,scaly=y,smooth=s
         3. cap-color:                brown=n,buff=b,cinnamon=c,gray=g,green=r, pink=p,purple=u,red=e,white=w,yellow=y
         4. bruises?:                 bruises=t,no=f
         5. odor:                     almond=a,anise=l,creosote=c,fishy=y,foul=f, musty=m,none=n,pungent=p,spicy=s
         6. gill-attachment:          attached=a,descending=d,free=f,notched=n
         7. gill-spacing:             close=c,crowded=w,distant=d
         8. gill-size:                broad=b,narrow=n
         9. gill-color:               black=k,brown=n,buff=b,chocolate=h,gray=g, green=r,orange=o,pink=p,purple=u,red=e,white=w,yellow=y
        10. stalk-shape:              enlarging=e,tapering=t
        11. stalk-root:               bulbous=b,club=c,cup=u,equal=e, rhizomorphs=z,rooted=r,missing=?
        12. stalk-surface-above-ring: fibrous=f,scaly=y,silky=k,smooth=s
        13. stalk-surface-below-ring: fibrous=f,scaly=y,silky=k,smooth=s
        14. stalk-color-above-ring:   brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
        15. stalk-color-below-ring:   brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
        16. veil-type:                partial=p,universal=u
        17. veil-color:               brown=n,orange=o,white=w,yellow=y
        18. ring-number:              none=n,one=o,two=t
        19. ring-type:                cobwebby=c,evanescent=e,flaring=f,large=l, none=n,pendant=p,sheathing=s,zone=z
        20. spore-print-color:        black=k,brown=n,buff=b,chocolate=h,green=r, orange=o,purple=u,white=w,yellow=y
        21. population:               abundant=a,clustered=c,numerous=n, scattered=s,several=v,solitary=y
        22. habitat:                  grasses=g,leaves=l,meadows=m,paths=p, urban=u,waste=w,woods=d

- Missing Attribute Values: 2480 of them (denoted by "?"), all for  attribute #11.

- Class Distribution: 
        --    edible: 4208 (51.8%)
        -- poisonous: 3916 (48.2%)
        --     total: 8124 instances

```python
#load the data set
data=pd.read_csv('./mushrooms.csv')
data.columns=columns=['class', 'cap-shape','cap-surface','cap-color','bruises?','odor','gill-attachment','gill-spacing','gill-size','gill-color','stalk-shape','stalk-root','stalk-surface-above-ring','stalk-surface-below-ring','stalk-color-above-ring','stalk-color-below-ring','veil-type','veil-color','ring-number','ring-type','spore-print-color','population','habitat']
print data.describe()
```

Output:
```
class cap-shape cap-surface cap-color bruises?  odor gill-attachment  \
count   8124      8124        8124      8124     8124  8124            8124   
unique     2         6           4        10        2     9               2   
top        e         x           y         n        f     n               f   
freq    4208      3656        3244      2284     4748  3528            7914   

       gill-spacing gill-size gill-color   ...   stalk-surface-below-ring  \
count          8124      8124       8124   ...                       8124   
unique            2         2         12   ...                          4   
top               c         b          b   ...                          s   
freq           6812      5612       1728   ...                       4936   

       stalk-color-above-ring stalk-color-below-ring veil-type veil-color  \
count                    8124                   8124      8124       8124   
unique                      9                      9         1         
```

```python
data.tail()
```

Output:
```
class cap-shape cap-surface cap-color bruises? odor gill-attachment  \
8119     e         k           s         n        f    n               a   
8120     e         x           s         n        f    n               a   
8121     e         f           s         n        f    n               a   
8122     p         k           y         n        f    y               f   
8123     e         x           s         n        f    n               a   

     gill-spacing gill-size gill-color   ...   stalk-surface-below-ring  \
8119            c         b          y   ...                          s   
8120            c         b          y   ...                          s   
8121            c         b          n   ...                          s   
8122            c         n          b   ...                          k   
8123            c         b          y   ...                          s   

     stalk-color-above-ring stalk-color-below-ring veil-type veil-color  \
8119                  
```

```python
# Get your data in matrix
X=np.asarray(data.ix[:,1:].dropna())
Y=np.asarray(data.ix[:,0].dropna())
print " Data Set Dimensions=", X.shape, " True Class labels dimensions", Y.shape   
```

Output:
```
Data Set Dimensions= (8124, 22)  True Class labels dimensions (8124,)
/home/arbaz/anaconda2/lib/python2.7/site-packages/ipykernel_launcher.py:2: DeprecationWarning: 
.ix is deprecated. Please use
.loc for label based indexing or
.iloc for positional indexing

See the documentation here:
http://pandas.pydata.org/pandas-docs/stable/indexing.html#ix-indexer-is-deprecated
```

```python
# Split your data into training and test-set... 
# see the documentation of split_data in tools for further information...
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

print " Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape   
print " Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytrain.shape   
```

Output:
```
Training Data Set Dimensions= (5687, 22) Training True Class labels dimensions (5687,)
 Test Data Set Dimensions= (2437, 22) Test True Class labels dimensions (5687,)
```

```python
# Lets train a Decision Tree Classifier two features
%pdb off
feat=np.arange(2)
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)
# print Xtrain
```

Output:
```
Automatic pdb calling has been turned OFF
size is  5687
size is  5346
size is  1565
size is  46
size is  1519
size is  3781
size is  3248
size is  1928
size is  1320
size is  533
size is  294
size is  239
size is  341
size is  111
size is  230
```

```python
print dt
```

Output:
```
---------------------------------------------------
 A Decision Tree With Depth=5
        C(class=e,Purity=0.934782608696)
    I(Fidx=1,Score=0.442749653452,Split=['y' 'g'])
        C(class=e,Purity=0.837837837838)
I(Fidx=0,Score=0.96742582674,Split=['x' 'c' 'k' 'f'])
                C(class=p,Purity=0.903765690377)
            I(Fidx=1,Score=0.697262115546,Split=['s'])
                C(class=p,Purity=0.690476190476)
        I(Fidx=0,Score=0.960619647823,Split=['x' 'f'])
                C(class=p,Purity=0.577272727273)
            I(Fidx=1,Score=0.992608681266,Split=['y'])
                C(class=p,Purity=0.514522821577)
    I(Fidx=1,Score=0.967072196056,Split=['f'])
            C(class=e,Purity=0.649769585253)
        I(Fidx=0,Score=0.917035960471,Split=['k'])
            C(class=e,Purity=0.934782608696)
---------------------------------------------------
```

```python
pclasses=dt.predict(Xtest[:,feat])
# print pclasses
#Lets see how good we are doing, by finding the accuracy on the test set..
print np.sum(pclasses==Ytest)
print "Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0])
```

Output:
```
1517
Accuracy =  0.622486663931
```

<!-- grade_id: acc, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)
feat=np.arange(2)
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)


pclasses=dt.predict(Xtest[:,feat])
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])

assert_greater_equal(acc, 0.60)
```

Output:
```
size is  5687
size is  5353
size is  1587
size is  44
c is  e
size is  1543
size is  3766
size is  3236
size is  1904
size is  1332
size is  530
size is  288
size is  242
size is  334
size is  106
size is  228
```

```python
from nose.tools import assert_greater_equal
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

dt=DecisionTree(0.95,5,10)
dt.train(Xtrain,Ytrain)

pclasses=dt.predict(Xtest)
print pclasses
```

Output:
```
size is  5687
size is  3033
c is  e
size is  2654
c is  p
['e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', 'e', '
```

<!-- grade_id: acc_all, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

dt=DecisionTree(0.95,5,10)
dt.train(Xtrain,Ytrain)

pclasses=dt.predict(Xtest)
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])

assert_greater_equal(acc, 0.90)
```

Output:
```
size is  5687
size is  3037
c is  e
size is  2650
c is  p
```

# Lets Train on all the features and for both the classes....

```python
# Split your data into training and test-set... 
# see the documentation of split_data in tools for further information...
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

print " Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape   
print " Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytrain.shape   
```

Output:
```
Training Data Set Dimensions= (5687, 22) Training True Class labels dimensions (5687,)
 Test Data Set Dimensions= (2437, 22) Test True Class labels dimensions (5687,)
```

```python
feat=arange(22)#[0, 1, 2, 3]
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)
pclasses=dt.predict(Xtest[:,feat])
#Lets see how good we are doing, by finding the accuracy on the test set..
print np.sum(pclasses==Ytest)
print "Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0])
```

Output:
```
size is  5687
size is  3028
c is  e
size is  2659
c is  p
2399
Accuracy =  0.984407057858
```

```python
print dt
```

Output:
```
---------------------------------------------------
 A Decision Tree With Depth=2
    C(class=p,Purity=1.0)
I(Fidx=4,Score=0.0955909088803,Split=['a' 'l' 'n'])
    C(class=e,Purity=0.972919418758)
---------------------------------------------------
```

What can you conclude ?
====================
Please write your observation....



```python

```

# Cross-Validation

Until now we have been splitting the dataset into a training and test set rather randomly and were reporting a rather artifical performance. Now we are going to test our system exhaustively by making use of k-fold [cross validation](http://en.wikipedia.org/wiki/Cross-validation_%28statistics%29). 

Now go and tune your hyper-parameters (purity, exthreshold) to opitmize the performance.

```python
# Now lets cross validate for best paramters, and test the result...
# We will be training four different models on four different partitions of data set and 
# then will be reporting the mean accuracy of the four classifiers.

nfolds=4 # lets use four folds..
folds=t.generate_folds(X,Y,nfolds)
features=arange(22)#[0, 1, 2, 3] # features to use for our system
#now lets train and test on these folds...

dt = DecisionTree(0.95,5)
totacc = []
for f in range(nfolds):

    traindata=folds[f][0]
    trainlabels=folds[f][1]

    #Lets first train the classifier
#         nb.train(traindata,trainlabels)
    dt.train(traindata[:,features],trainlabels)
#                 knn.train(traindata[:,features],trainlabels)

    testdata=folds[f][2]
    testlabels=folds[f][3]

    #Lets test the classifier
#         pclasses= nb.test(testdata)
    pclasses=dt.predict(testdata[:,features])

    #print pclasses
    acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
    print "[Info] Fold {} Accuracy = {}".format(f+1, acc)    
    totacc.append(acc)
#foldacc.append(np.mean(np.array(totacc)))
# params[i][j] = np.mean(totacc)
# print  "purity = ",i,"Nexample-threshold = ",j,"Mean Accuracy = ",params[i][j]


```

Output:
```
Generating CV data for 2 classes
size is  6093
size is  3251
c is  e
size is  2842
c is  p
[Info] Fold 1 Accuracy = 0.987690792713
size is  6093
size is  3235
c is  e
size is  2858
c is  p
[Info] Fold 2 Accuracy = 0.979812900049
size is  6093
size is  3251
c is  e
size is  2842
c is  p
[Info] Fold 3 Accuracy = 0.987690792713
size is  6093
size is  3247
c is  e
size is  2846
c is  p
[Info] Fold 4 Accuracy = 0.985721319547
```

```python
# Now lets cross validate for best paramters, and test the result...
# We will be training four different models on four different partitions of data set and 
# then will be reporting the mean accuracy of the four classifiers.

nfolds=4 # lets use four folds..
folds=t.generate_folds(X,Y,nfolds)
features=[0,1, 2, 3] # features to use for our system
#now lets train and test on these folds...

#Lets perform the grid search...
purity=np.linspace(0.85,0.97,13)
nexamp=np.linspace(5,25,21)  

params=np.zeros((len(purity),len(nexamp)))

for i in range(len(purity)):
    for j in range(len(nexamp)):
        dt = DecisionTree(purity[i],nexamp[j])
        totacc = []
        for f in range(nfolds):

            traindata=folds[f][0]
            trainlabels=folds[f][1]

            #Lets first train the classifier
    #         nb.train(traindata,trainlabels)
            dt.train(traindata[:,features],trainlabels)
#                 knn.train(traindata[:,features],trainlabels)

            testdata=folds[f][2]
            testlabels=folds[f][3]

            #Lets test the classifier
    #         pclasses= nb.test(testdata)
            pclasses=dt.predict(testdata[:,features])

            #print pclasses
            acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
            print "[Info] Fold {} Accuracy = {}".format(f+1, acc)    
            totacc.append(acc)
        #foldacc.append(np.mean(np.array(totacc)))
        params[i][j] = np.mean(totacc)
        print  "purity = ",i,"Nexample-threshold = ",j,"Mean Accuracy = ",params[i][j]
                   

```

Output:
```
Generating CV data for 2 classes
size is  6093
size is  5727
size is  1667
size is  724
c is  e
size is  943
size is  642
size is  16
c is  e
size is  626
size is  301
c is  p
size is  4060
size is  3488
size is  2693
size is  26
c is  e
size is  2667
size is  463
size is  1
c is  p
size is  462
size is  2204
size is  440
size is  1764
size is  671
size is  1093
size is  795
size is  572
size is  64
c is  e
size is  508
c is  p
size is  366
c is  e
[Info] Fold 1 Accuracy = 0.679960610537
size is  6093
size is  1744
size is  1448
size is  667
size is  612
size is  55
c is  e
size is  781
c is  e
size is  296
c is  p
size is  4349
size is  291
c is  e
size is  4058
size is  3473
size is  1229
size is  108
size is  1121
size is  2244
size is  22
c is  e
size is  2222
size is  464
size is  1
c is  p
size is  463
size is  1758
size is  677
size is  1081
size is  585
size is  69
c is  e
size is  516
c is  p
[Info] Fold 2 Accuracy = 0.66765140325
size is  6093
size is  1748
size is  788
c is 
```

```python

```